# Homework 3: Neural Networks


This assignment is due on Moodle by **11:59pm on Friday October 19**. 
Your solutions to theoretical questions should be done in Markdown/MathJax directly below the associated question.
Your solutions to computational questions should include any specified Python code and results 
as well as written commentary on your conclusions.
Remember that you are encouraged to discuss the problems with your instructors and classmates, 
but **you must write all code and solutions on your own**. For a refresher on the course **Collaboration Policy** click [here](https://github.com/BoulderDS/CSCI-4622-Machine-Learning-18fa/blob/master/info/syllabus.md#collaboration-policy).

**NOTES**: 

- Do **NOT** load or use any Python packages that are not available in Anaconda 3.6. 
- Some problems with code may be autograded.  If we provide a function API **do not** change it.  If we do not provide a function API then you're free to structure your code however you like. 
- Submit only this Jupyter notebook to Moodle.  Do not compress it using tar, rar, zip, etc. 

** Acknowledgment ** : Chris Ketelsen


Name: 

```python
import math
import pickle
import gzip
import numpy as np
import pandas
import matplotlib.pylab as plt
%matplotlib inline
```

[25 points] Problem 1 - Single-Layer and Multilayer Perceptron Learning
---

**Part A** : Consider learning the following concepts with either a single-layer or multilayer perceptron where all hidden and output neurons utilize *indicator* activation functions. For each of the following concepts, state whether the concept can be learned by a single-layer perceptron. Briefly justify your response by providing weights and biases as applicable:

i. $~ \texttt{ NOT } x_1$

ii. $~~x_1 \texttt{ NAND } x_2$

iii. $~~x_1 \texttt{ XNOR } x_2$ (output 1 when $x_1 = x_2$ and 0 otherwise)

<!-- grade_id: cell-5504bb80827fe61b, points: 15, grade: true, solution: true -->
YOUR ANSWER HERE

**Part B** : Determine an architecture and specific values of the weights and biases in a single-layer or multilayer perceptron with *indicator* activation functions that can learn $x_1 \texttt{ XNOR } x_2$. Describe your architecture and state your weight matrices and bias vectors in Markdown below. Then demonstrate that your solution is correct by implementing forward propagation for your network in Python and showing that it correctly produces the correct boolean output values for each of the four possible combinations of $x_1$ and $x_2$.

<!-- grade_id: cell-fd1e475a5ef92def, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-330996afb5a81650, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

[25 points] Problem 2 - Back propagation
---


In this problem you'll gain some intuition about why training deep neural networks can be very time consuming.  Consider training the chain-like neural network seen below: 

![chain-like nn](figs/chain_net.png)

Note that this network has three weights $W^1, W^2, W^3$ and three biases $b^1, b^2,$ and $b^3$ (for this problem you can think of each parameter as a single value or as a $1 \times 1$ matrix). Suppose that each hidden and output neuron is equipped with a sigmoid activation function and the loss function is given by 

$$
\ell(y, a^4) = \frac{1}{2}(y - a^4)^2  
$$

where $a^4$ is the value of the activation at the output neuron and $y \in \{0,1\}$ is the true label associated with the training example. 

**Part A**: Suppose each of the weights is initialized to $W^k = 1.0$ and each bias is initialized to $b^k = -0.5$.  Use forward propagation to find the activities and activations associated with each hidden and output neuron for the training example $(x, y) = (0.5,0)$. Show your work. 

<!-- grade_id: cell-6512c42fc5e9ce1b, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

**Part B**: Use Back-Propagation to compute the weight and bias derivatives $\partial \ell / \partial W^k$ and $\partial \ell / \partial b^k$ for $k=1, 2, 3$.  Show all work. 

<!-- grade_id: cell-6a3b895fa888e925, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

**PART C** Implement following activation functions:
* Relu
* Sigmoid
* softmax

<!-- grade_id: cell-3b7e3adaffe2c2ee, points: 0, grade: false, solution: true -->
```python
import math

def relu(x):
    # YOUR CODE HERE
    raise NotImplementedError()

def sigmoid(x):
    # YOUR CODE HERE
    raise NotImplementedError()

def soft_max(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0bded3752d3226b9, points: 5, grade: true, solution: false -->
```python
assert relu(5) == 5
assert relu(-5) == 0
assert relu(0) == 0
assert relu(5) == 5

assert sigmoid(0.458) == 0.61253961344091512
assert sigmoid(2) == 0.8807970779778823
res = soft_max([1,2,4])
temp = [0.04201007, 0.1141952 , 0.84379473]
for i in range(len(temp)):
    assert res[i] - temp[i] < 1e-5
```

**PART D** Implement the following Loss functions:
* mean squared error
* mean absolute error
* hinge

<!-- grade_id: cell-84456d343bfeca31, points: 0, grade: false, solution: true -->
```python
def mean_squared_error(yhat,y):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def mean_absolute_error(yhat,y):
    # YOUR CODE HERE
    raise NotImplementedError()

def hinge(yhat,y):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-550e920d814cb6d3, points: 5, grade: true, solution: false -->
```python
y_true = np.array([3, -0.5, 2, 7])
y_pred = np.array([2.5, 0.0, 2, 8])
assert mean_squared_error(y_pred,y_true) == 0.375
assert mean_absolute_error(y_pred,y_true) == 0.5
assert hinge(y_pred,y_true) == 0.25
```

**PART E** Explain the vanishing gradient problem, when would you observe this?

<!-- grade_id: cell-59ccf91056b5d8bb, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

[25 Points] Problem 3 - Build a feedForward neural network
---

In this problem you'll implement a general feed-forward neural network class that utilizes sigmoid activation functions. Your tasks will be to implement forward propagation, prediction, back propagation, and a general train routine to learn the weights in your network via stochastic gradient descent.

The skeleton for the network class is below. Note that this class is almost identical to the one you worked with in the hands-On neural network in-class notebook, so you should look there to remind yourself of the details. Scroll down to find more information about your tasks as well as unit tests.

<!-- grade_id: cell-b8abc0ac570aac74, points: 0, grade: false, solution: true -->
```python
import argparse
import numpy as np
import pickle
import gzip
import matplotlib.pyplot as plt


class Network:
    def __init__(self, sizes):
        self.L = len(sizes)
        self.sizes = sizes
        self.biases = [np.random.randn(n, 1) for n in self.sizes[1:]]
        self.weights = [np.random.randn(n, m) for (
            m, n) in zip(self.sizes[:-1], self.sizes[1:])]

    def g(self, z):
        """
        activation function
        """
        return sigmoid(z)

    def g_prime(self, z):
        """
        derivative of activation function
        """
        return sigmoid_prime(z)

    def forward_prop(self, a):
        """
        memory aware forward propagation for testing
        only.  back_prop implements it's own forward_prop
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def grad_cost(self, a, y):
        """
        gradient of cost function
        Assumes C(a,y) = (a-y)^2/2
        """
        return (a - y)

    def SGD_train(self, train, epochs, eta, lam=0.0, verbose=True, test=None):
        """
        SGD for training parameters
        epochs is the number of epocs to run
        eta is the learning rate
        lam is the regularization parameter
        If verbose is set will print progressive accuracy updates
        If test set is provided, routine will print accuracy on test set as learning evolves
        """
        n_train = len(train)
        epoch_accuracy_list = []  # omit
        hidden_acc_train = 0  # omit
        hidden_acc_test = 0  # omit
        # YOUR CODE HERE
        raise NotImplementedError()
            if verbose:
                if epoch == 0 or (epoch + 1) % 15 == 0:
                    acc_train = self.evaluate(train)
                    if test is not None:
                        acc_test = self.evaluate(test)
                        # epoch_accuracy_list.append((acc_test,acc_train))    #omit
                        print("Epoch {:4d}: Train {:10.5f}, Test {:10.5f}".format(
                            epoch+1, acc_train, acc_test))
                    else:
                        print("Epoch {:4d}: Train {:10.5f}".format(
                            epoch+1, acc_train))

    def back_prop(self, x, y):
        """
        Back propagation for derivatives of C wrt parameters
        """
        db_list = [np.zeros(b.shape) for b in self.biases]
        dW_list = [np.zeros(W.shape) for W in self.weights]

        a = x
        a_list = [a]
        z_list = [np.zeros(a.shape)]  # Pad with throwaway so indices match

        for W, b in zip(self.weights, self.biases):
            z = np.dot(W, a) + b
            z_list.append(z)
            a = self.g(z)
            a_list.append(a)

        # Back propagate deltas to compute derivatives
        # The following list gives hints on how to do it
        # calculating delta (Error) for the output layer
        # for the appropriate layers compute db_list[ell], dW_list[ell], delta

        # YOUR CODE HERE
        raise NotImplementedError()
        return (dW_list, db_list)

    def evaluate(self, test):
        """
        Evaluate current model on labeled test data
        """
        ctr = 0
        for x, y in test:
            yhat = self.forward_prop(x)
            ctr += np.argmax(yhat) == np.argmax(y)
        return float(ctr) / float(len(test))


def sigmoid(z, threshold=20):
    z = np.clip(z, -threshold, threshold)
    return 1.0/(1.0+np.exp(-z))


def sigmoid_prime(z):
    return sigmoid(z) * (1.0 - sigmoid(z))


def mnist_digit_show(flatimage, outname=None):

    import matplotlib.pyplot as plt

    image = np.reshape(flatimage, (-1, 14))

    plt.matshow(image, cmap=plt.cm.binary)
    plt.xticks([])
    plt.yticks([])
    if outname:
        plt.savefig(outname)
    else:
        plt.show()
```

**PART A** Implement *SGD_train, back_prop,forward_prop*. Use the following test cases to verify if the code is working as expected.

<!-- grade_id: cell-7632a78793a2588e, points: 15, grade: true, solution: false -->
```python
%run -i tests/tests.py
```

** PART B **

Run the above Network on MNIST Dataset and report the following (feel free to experiment with different learning rates).
* Define the hidden layer dimension appropriate for this dataset and report the accuracy for 200 epochs.
* Explain the effect of hidden dimension with appropriate plots (check training accuracy and test accuracy). 
* Explain the effect of number of epochs on MNIST Data.

```python
location = './data/tinyMNIST.pkl.gz'
f = gzip.open(location, 'rb')
u = pickle._Unpickler(f)
u.encoding = 'latin1'
train, test = u.load()
input_dimensions = len(train[0][0])
output_dimensions = len(train[0][1])
hidden_layer_dimensions = 10
print('Number of Input Features: ', input_dimensions)
print('Number of Output classes: ', output_dimensions)
print('\nHidden Layer Dimensions: ', hidden_layer_dimensions)
nn = Network([input_dimensions, hidden_layer_dimensions, output_dimensions])
nn.SGD_train(train, epochs=200, eta=0.1, lam=0.0001, verbose=True, test=test)
```

#### [Extra Credit (10 points)] Implement *Dropout layer* and explain the impact on performance
---

<!-- grade_id: cell-f7345360536d4c03, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

[25 Points] Problem 4 - Implement RNN Network to classify whether text is spam or ham 
---

Dataset is obtained from UCI Machine Learning repository consisting of SMS tagged messages (being ham (legitimate) or spam) that have been collected for SMS Spam research.


Use [Keras](https://keras.io/) to implement a classifier (you need to install Keras). Update the below snippet to build a Sequential model with an embedding layer, and an LSTM layer followed by a dense layer. This question allows you to get familiar with popular deep learning toolkits and the solution only has a few lines. In practice, there is no need to reinvent the wheels.


Learn more about RNN : https://colah.github.io/posts/2015-08-Understanding-LSTMs/

<!-- grade_id: cell-e5faa2bc69b4d8cf, points: 0, grade: false, solution: true -->
```python
from keras.datasets import imdb
from keras.models import Sequential
from keras.layers import Dense
from keras.layers import LSTM
from keras.layers.embeddings import Embedding
from keras.preprocessing import sequence
from keras.preprocessing.text import Tokenizer


class RNN:
    '''
    RNN classifier
    '''

    def __init__(self, train_x, train_y, test_x, test_y, dict_size=5000,
                 example_length=150, embedding_length=32, epoches=5, batch_size=128):
        '''
        initialize RNN model
        :param train_x: training data
        :param train_y: training label
        :param test_x: test data
        :param test_y: test label
        :param epoches: number of ephoches to run
        :param batch_size: batch size in training
        :param embedding_length: size of word embedding
        :param example_length: length of examples
        '''
        self.batch_size = batch_size
        self.epoches = epoches
        self.example_len = example_length
        self.dict_size = dict_size
        self.embedding_len = embedding_length

        # preprocess training data
        tok = Tokenizer(num_words=dict_size)
        tok.fit_on_texts(train_x)
        sequences = tok.texts_to_sequences(train_x)
        self.train_x = sequence.pad_sequences(
            sequences, maxlen=self.example_len)
        sequences = tok.texts_to_sequences(test_x)
        self.test_x = sequence.pad_sequences(
            sequences, maxlen=self.example_len)

        self.train_y = train_y
        self.test_y = test_y

        # TODO: build model with Embedding, LSTM and dense layers.
        # refer to Sequence classification with LSTM : https://keras.io/getting-started/sequential-model-guide/#examples
        # Documentation for LSTM layer in : https://keras.io/layers/recurrent/#lstm
        self.model = Sequential()
        # YOUR CODE HERE
        raise NotImplementedError()
        self.model.compile(loss='binary_crossentropy',
                           optimizer='adam', metrics=['accuracy'])

    def train(self, verbose=0):
        '''
        fit in data and train model : refer fit method in https://keras.io/models/model/
        make sure you use batchsize and epochs appropriately.
        :return:None
        '''
        # TODO: fit in data to train your model
        # YOUR CODE HERE
        raise NotImplementedError()

    def evaluate(self):
        '''

        evaluate trained model : Please refer evaluate in https://keras.io/models/model/
        :return: [loss,accuracy]
        '''
        # YOUR CODE HERE
        raise NotImplementedError()
```

Complete the following functions *init, train and evaluate functions* and report the accuracy

```python
import keras
assert keras.__version__ == '2.2.2'
```

```python
import pickle
def load_data(location):
    return pickle.load(open(location,'rb'))
```

```python
train_x, test_x, train_y, test_y = load_data('./data/spam_data.pkl')
rnn = RNN(train_x, train_y, test_x, test_y, epoches=5)
rnn.train(verbose=1)
accuracy = rnn.evaluate()
print('Accuracy for LSTM: ', accuracy[1])
```

Accuracy after 5 epochs.

```python

```

* Change the embedding length and observe the impact on test and train accuracy.

* Explain the impact of embedding length in LSTM Model by providing plots of accuracy vs embedding length.

<!-- grade_id: cell-4c9a9edc927e69ee, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

```python

```